```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

# Lab 7 – Regular Expressions

## DSC 80, Fall 2025

### Due Date: Monday, November 17th at 11:59 PM

## Instructions

Welcome to the seventh DSC 80 lab this quarter!

Much like in DSC 10, this Jupyter Notebook contains the statements of the problems and provides code and Markdown cells to display your answers to the problems. Unlike DSC 10, the notebook is *only* for displaying a readable version of your final answers. The coding will be done in an accompanying `lab.py` file that is imported into the current notebook, and **you will only submit that `lab.py` file**, not this notebook!

Some additional guidelines:
- **Unlike in DSC 10, labs will have both public tests and hidden tests.** The bulk of your grade will come from your scores on hidden tests, which you will only see on Gradescope after the assignment deadline.
- **Do not change the function names in the `lab.py` file!** The functions in the `lab.py` file are how your assignment is graded, and they are graded by their name. If you changed something you weren't supposed to, you can find the original code in the [course GitHub repository](https://github.com/dsc-courses/dsc80-2025-sp).
- Notebooks are nice for testing and experimenting with different implementations before designing your function in your `lab.py` file. You can write code here, but make sure that all of your real work is in the `lab.py` file, since that's all you're submitting.
- You are encouraged to write your own additional helper functions to solve the lab, as long as they also end up in `lab.py`.

**To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.** More details on its usage are given at the bottom of this notebook.

**Importing code from `lab.py`**:

* Below, we import the `.py` file that's contained in the same directory as this notebook.
* We use the `autoreload` notebook extension to make changes to our `lab.py` file immediately available in our notebook. Without this extension, we would need to restart the notebook kernel to see any changes to `lab.py` in the notebook.
    - `autoreload` is necessary because, upon import, `lab.py` is compiled to bytecode (in the directory `__pycache__`). Subsequent imports of `lab` merely import the existing compiled python.

```python
%load_ext autoreload
%autoreload 2
```

Output:
```
The autoreload extension is already loaded. To reload it, use:
  %reload_ext autoreload
```

```python
from lab import *
```

```python
import pandas as pd
import numpy as np
import os
import re
```

## Question 1 – Practice with Regular Expressions 🛠
<div class="alert alert-block alert-warning">
<b>Note</b>: The functions in this question all have doctests in their docstrings. The doctests constitute the public tests; we will still run hidden tests on each of your functions. <b>Please don't change any of the docstrings!</b>
</div>

Regular expressions can be tricky, and the best way to gain familiarity with them is through lots of practice. In this question, you will work through ten exercises, each of which requires you to write a regular expression that matches strings that satisfy certain criteria. Make sure to take a close look at the doctests for each function in `lab.py`, as they provide useful guidance for the types of strings you should and shouldn't match.

***Notes:*** 
- Make sure to refer to the [Regular Expression Resources](https://dsc80.com/resources/#regular-expressions) posted on the course website. In particular, we recommend having [regex101.com](https://regex101.com/) open while working, along with the [cheat sheet](https://dsc80.com/resources/other/berkeley-regex-reference.pdf).

- Each exercise has a star rating, between 1 (⭐️) and 3 (⭐️⭐️⭐️) stars, indicating its difficulty level (1 being the easiest, 3 being the hardest). If you are spending lots of time on 1-star exercises, take a close look at the syntax from lecture, as there is probably an easier way of writing the necessary pattern!

<br>

### Exercise 1 (⭐️)

Write a regular expression that matches strings that have `'['` as the third character and `']'` as the sixth character.

<br>

### Exercise 2 (⭐️)

Write a regular expression that matches strings that are phone numbers that start with `'(858)'` and follow the format `'(xxx) xxx-xxxx'` (`'x'` represents a digit).

***Note:*** There is a space between `'(xxx)'` and `'xxx-xxxx'`.

<br>

### Exercise 3 (⭐️)

Write a regular expression that matches strings that:
- are between 6 and 10 characters long (inclusive),
- contain only alphanumeric characters, whitespace and `'?'`, and
- end with `'?'`.

<br>

### Exercise 4 (⭐️⭐️)

Write a regular expression that matches strings with exactly two `'$'`, one of which is at the start of the string, such that:
- the characters between the two `'$'` can be anything (including nothing) except the lowercase letters `'a'`, `'b'`, and `'c'`, (and `'$'`), and
- the characters after the second `'$'` can only be the **lowercase or uppercase** letters `'a'`/`'A'`, `'b'`/`'B'`, and `'c'`/`'C'`, with every `'a'`/`'A'` before every `'b'`/`'B'`, and every `'b'`/`'B'` before every `'c'`/`'C'`. There must be at least one `'a'` or `'A'`, at least one `'b'` or `'B'`, and at least one `'c'` or `'C'`.
    

<br>

### Exercise 5 (⭐️)
Write a regular expression that matches strings that represent valid Python file names, including the extension. 

***Note:*** For simplicity, assume that file names only contain letters, numbers, and underscores (`'_'`).

<br>

### Exercise 6 (⭐️)
Write a regular expression that matches strings that:
- are made up of only lowercase letters and exactly one underscore (`'_'`), and
- have at least one lowercase letter on both sides of the underscore.

<br>

### Exercise 7 (⭐️)
Write a regular expression that matches strings that start with and end with an underscore (`'_'`).

<br>

### Exercise 8 (⭐️)

Apple serial numbers are strings of length 1 or more that are made up of any characters, other than
- the uppercase letter `'O'`, 
- the lowercase letter `'i`', and 
- the number `'1'`.

Write a regular expression that matches strings that are valid Apple serial numbers.

<br>

### Exercise 9 (⭐️⭐️)

ID numbers are formatted as `'SC-NN-CCC-NNNN'`, where 
- SC represents state code in uppercase (e.g. `'CA'`),
- NN represents a number with 2 digits (e.g. `'98'`),
- CCC represents a three letter city code in uppercase (e.g. `'SAN'`), and
- NNNN represents a number with 4 digits (e.g. `'1024'`).

Write a regular expression that matches strings that are ID numbers corresponding to the cities of `'SAN'` or `'LAX'`, or the state of `'NY'`. Assume that there is only one city named `'SAN'`, only one city named `'LAX'`, and they are supposed to be in the `'CA'` state.

<br>

### Exercise 10 (⭐️⭐️⭐️)

Write a function named `match_10` that takes in a string and:
- converts the string to lowercase,
- removes all non-alphanumeric characters (i.e. removes everything that is not in the `\w` character class), and the letter `'a'`, and
- returns a list of every **non-overlapping** three-character substring in the remaining string, starting from the beginning of the string.
   
For instance, consider the following doctest:

```py
>>> match_10('Ab..DEF')
['bde']
```

Here's how `match_10` should process `'Ab..DEF'`:

1. Convert to lowercase: `'ab..def'`.
2. Remove non-alphanumeric characters and the letter `'a'`: `'bdef'`.
3. Starting from the beginning of the string, there is only a single non-overlapping three character substring: `'bde'`. Hence, we return `['bde']`.

***Note:*** Perform your operations in the exact order described above, otherwise your code may not pass all the tests. Don't use a `for`-loop.

```python
def match_1(string):
    """
    DO NOT EDIT THE DOCSTRING!
    >>> match_1("abcde]")
    False
    >>> match_1("ab[cde")
    False
    >>> match_1("a[cd]")
    False
    >>> match_1("ab[cd]")
    True
    >>> match_1("1ab[cd]")
    False
    >>> match_1("ab[cd]ef")
    True
    >>> match_1("1b[#d] _")
    True
    """
    pattern = r'^.{2}\[.{2}\]'
    
    # Do not edit following code
    prog = re.compile(pattern)
    return prog.search(string) is not None

print(match_1("abcde]"))
print(match_1("ab[cde"))
print(match_1("a[cd]"))
print(match_1("ab[cd]"))
print(match_1("1ab[cd]"))
print(match_1("ab[cd]ef"))
print(match_1("1b[#d] _"))
```

Output:
```
False
False
False
True
False
True
True
```

```python
def match_2(string):
    """
    DO NOT EDIT THE DOCSTRING!
    >>> match_2("(123) 456-7890")
    False
    >>> match_2("858-456-7890")
    False
    >>> match_2("(858)45-7890")
    False
    >>> match_2("(858) 456-7890")
    True
    >>> match_2("(858)456-789")
    False
    >>> match_2("(858)456-7890")
    False
    >>> match_2("a(858) 456-7890")
    False
    >>> match_2("(858) 456-7890b")
    False
    """
    pattern = r'^\(858\) \d{3}\-\d{4}$'

    # Do not edit following code
    prog = re.compile(pattern)
    return prog.search(string) is not None

# print(match_2("(123) 456-7890"))
# print(match_2("858-456-7890"))
# print(match_2("(858)45-7890"))
# print(match_2("(858) 456-7890"))
# print(match_2("(858)456-789"))
# print(match_2("(858)456-7890"))
# print(match_2("a(858) 456-7890"))
# print(match_2("(858) 456-7890b"))
```

```python
# ### Exercise 3 (⭐️)

# Write a regular expression that matches strings that:
# - are between 6 and 10 characters long (inclusive),
# - contain only alphanumeric characters, whitespace and `'?'`, and
# - end with `'?'`.

def match_3(string):
    """
    DO NOT EDIT THE DOCSTRING!
    >>> match_3("qwertsd?")
    True
    >>> match_3("qw?ertsd?")
    True
    >>> match_3("ab c?")
    False
    >>> match_3("ab   c ?")
    True
    >>> match_3(" asdfqwes ?")
    False
    >>> match_3(" adfqwes ?")
    True
    >>> match_3(" adf!qes ?")
    False
    >>> match_3(" adf!qe? ")
    False
    """
    pattern = r'^[a-zA-Z0-9\s?]{5,9}\?'

    # Do not edit following code
    prog = re.compile(pattern)
    return prog.search(string) is not None

# print(match_3("qwertsd?"))
# print(match_3("qw?ertsd?"))
# print(match_3("ab c?"))
# print(match_3("ab   c ?"))
# print(match_3(" asdfqwes ?"))
# print(match_3(" adfqwes ?"))
# print(match_3(" adf!qes ?"))
# print(match_3(" adf!qe? "))
```

```python
# ### Exercise 4 (⭐️⭐️)

# Write a regular expression that matches strings with exactly two `'$'`, one of which is at the start of the string, such that:
# - the characters between the two `'$'` can be anything (including nothing) except the lowercase letters `'a'`, `'b'`, and `'c'`, (and `'$'`), and
# - the characters after the second `'$'` can only be the **lowercase or uppercase** letters `'a'`/`'A'`, `'b'`/`'B'`, and `'c'`/`'C'`, with every `'a'`/`'A'` before every `'b'`/`'B'`, and every `'b'`/`'B'` before every `'c'`/`'C'`. There must be at least one `'a'` or `'A'`, at least one `'b'` or `'B'`, and at least one `'c'` or `'C'`.

def match_4(string):
    """
    DO NOT EDIT THE DOCSTRING!
    >>> match_4("$$AaaaaBbbbc")
    True
    >>> match_4("$!@#$aABc")
    True
    >>> match_4("$a$aABc")
    False
    >>> match_4("$iiuABc")
    False
    >>> match_4("123$$$Abc")
    False
    >>> match_4("$$Abc")
    True
    >>> match_4("$qw345t$AAAc")
    False
    >>> match_4("$s$Bca")
    False
    >>> match_4("$!@$")
    False
    """
    pattern = r'^\$[^abc$]*\$[aA]+[bB]+[cC]+$'

    # Do not edit following code
    prog = re.compile(pattern)
    return prog.search(string) is not None

# print(match_4("$$AaaaaBbbbc"))
# print(match_4("$!@#$aABc"))
# print(match_4("$a$aABc"))
# print(match_4("$iiuABc"))
# print(match_4("123$$$Abc"))
# print(match_4("$$Abc"))
# print(match_4("$qw345t$AAAc"))
# print(match_4("$s$Bca"))
# print(match_4("$!@$"))
```

```python
# ### Exercise 5 (⭐️)
# Write a regular expression that matches strings that represent valid Python file names, including the extension. 

# ***Note:*** For simplicity, assume that file names only contain letters, numbers, and underscores (`'_'`).

def match_5(string):
    """
    DO NOT EDIT THE DOCSTRING!
    >>> match_5("dsc80.py")
    True
    >>> match_5("dsc80py")
    False
    >>> match_5("dsc80..py")
    False
    >>> match_5("dsc80+.py")
    False
    """
    pattern = r'^[a-zA-Z0-9_]+\.py$'

    # Do not edit following code
    prog = re.compile(pattern)
    return prog.search(string) is not None

print(match_5("dsc80.py"))
print(match_5("dsc80py"))
print(match_5("dsc80..py"))
print(match_5("dsc80+.py"))
```

Output:
```
True
False
False
False
```

```python
# ### Exercise 6 (⭐️)
# Write a regular expression that matches strings that:
# - are made up of only lowercase letters and exactly one underscore (`'_'`), and
# - have at least one lowercase letter on both sides of the underscore.

def match_6(string):
    """
    DO NOT EDIT THE DOCSTRING!
    >>> match_6("aab_cbb_bc")
    False
    >>> match_6("aab_cbbbc")
    True
    >>> match_6("aab_Abbbc")
    False
    >>> match_6("abcdef")
    False
    >>> match_6("ABCDEF_ABCD")
    False
    """
    pattern = r'^[a-z]+_[a-z]+$'

    # Do not edit following code
    prog = re.compile(pattern)
    return prog.search(string) is not None

print(match_6("aab_cbb_bc"))
print(match_6("aab_cbbbc"))
print(match_6("aab_Abbbc"))
print(match_6("abcdef"))
print(match_6("ABCDEF_ABCD"))
```

Output:
```
False
True
False
False
False
```

```python
# ### Exercise 7 (⭐️)
# Write a regular expression that matches strings that start with and end with an underscore (`'_'`).

def match_7(string):
    """
    DO NOT EDIT THE DOCSTRING!
    >>> match_7("_abc_")
    True
    >>> match_7("abd")
    False
    >>> match_7("bcd")
    False
    >>> match_7("_ncde")
    False
    """
    pattern = r'^_.*_$'

    # Do not edit following code
    prog = re.compile(pattern)
    return prog.search(string) is not None

print(match_7("_abc_"))
print(match_7("abd"))
print(match_7("bcd"))
print(match_7("_ncde"))
```

Output:
```
True
False
False
False
```

```python
# ### Exercise 8 (⭐️)

# Apple serial numbers are strings of length 1 or more that are made up of any characters, other than
# - the uppercase letter `'O'`, 
# - the lowercase letter `'i`', and 
# - the number `'1'`.

# Write a regular expression that matches strings that are valid Apple serial numbers.

def match_8(string):
    """
    DO NOT EDIT THE DOCSTRING!
    >>> match_8("ASJDKLFK10ASDO")
    False
    >>> match_8("ASJDKLFK0ASDo!!!!!!! !!!!!!!!!")
    True
    >>> match_8("JKLSDNM01IDKSL")
    False
    >>> match_8("ASDKJLdsi0SKLl")
    False
    >>> match_8("ASDJKL9380JKAL")
    True
    """
    pattern = r'^[^Oi1]+$'

    # Do not edit following code
    prog = re.compile(pattern)
    return prog.search(string) is not None

print(match_8("ASJDKLFK10ASDO"))
print(match_8("ASJDKLFK0ASDo!!!!!!! !!!!!!!!!"))
print(match_8("JKLSDNM01IDKSL"))
print(match_8("ASDKJLdsi0SKLl"))
print(match_8("ASDJKL9380JKAL"))
```

Output:
```
False
True
False
False
True
```

```python
# ### Exercise 9 (⭐️⭐️)

# ID numbers are formatted as `'SC-NN-CCC-NNNN'`, where 
# - SC represents state code in uppercase (e.g. `'CA'`),
# - NN represents a number with 2 digits (e.g. `'98'`),
# - CCC represents a three letter city code in uppercase (e.g. `'SAN'`), and
# - NNNN represents a number with 4 digits (e.g. `'1024'`).

# Write a regular expression that matches strings that are ID numbers corresponding to the cities of `'SAN'` or `'LAX'`, or the state of `'NY'`. Assume that there is only one city named `'SAN'`, only one city named `'LAX'`, and they are supposed to be in the `'CA'` state.

def match_9(string):
    '''
    DO NOT EDIT THE DOCSTRING!
    >>> match_9('NY-32-NYC-1232')
    True
    >>> match_9('ca-23-SAN-1231')
    False
    >>> match_9('MA-36-BOS-5465')
    False
    >>> match_9('CA-56-LAX-7895')
    True
    >>> match_9('NY-32-LAX-0000') # If the state is NY, the city can be any 3 letter code, including LAX or SAN!
    True
    >>> match_9('TX-32-SAN-4491')
    False
    '''
    pattern = r'^(CA-\d{2}-(SAN|LAX)-\d{4})|(NY-\d{2}-[A-Z]{3}-\d{4})$'

    # Do not edit following code
    prog = re.compile(pattern)
    return prog.search(string) is not None

print(match_9('NY-32-NYC-1232'))
print(match_9('ca-23-SAN-1231'))
print(match_9('MA-36-BOS-5465'))
print(match_9('CA-56-LAX-7895'))
print(match_9('NY-32-LAX-0000')) # If the state is NY, the city can be any 3 letter code, including LAX or SAN!
print(match_9('TX-32-SAN-4491'))
```

Output:
```
True
False
False
True
True
False
```

```python
# ### Exercise 10 (⭐️⭐️⭐️)

# Write a function named `match_10` that takes in a string and:
# - converts the string to lowercase,
# - removes all non-alphanumeric characters (i.e. removes everything that is not in the `\w` character class), and the letter `'a'`, and
# - returns a list of every **non-overlapping** three-character substring in the remaining string, starting from the beginning of the string.
   
# For instance, consider the following doctest:

# ```py
# >>> match_10('Ab..DEF')
# ['bde']
# ```

# Here's how `match_10` should process `'Ab..DEF'`:

# 1. Convert to lowercase: `'ab..def'`.
# 2. Remove non-alphanumeric characters and the letter `'a'`: `'bdef'`.
# 3. Starting from the beginning of the string, there is only a single non-overlapping three character substring: `'bde'`. Hence, we return `['bde']`.

# ***Note:*** Perform your operations in the exact order described above, otherwise your code may not pass all the tests. Don't use a `for`-loop.

def match_10(string):
    '''
    DO NOT EDIT THE DOCSTRING!
    >>> match_10('ABCdef')
    ['bcd']
    >>> match_10(' DEFaabc !g ')
    ['def', 'bcg']
    >>> match_10('Come ti chiami?')
    ['com', 'eti', 'chi']
    >>> match_10('and')
    []
    >>> match_10('Ab..DEF')
    ['bde']
    
    '''
    string = string.lower()
    string = re.sub(r'[^\w]|a', "", string)
    string = re.findall(r'.{3}', string)
    return string

print(match_10('ABCdef'))
print(match_10(' DEFaabc !g '))
print(match_10('Come ti chiami?'))
print(match_10('and'))
print(match_10('Ab..DEF'))
```

Output:
```
['bcd']
['def', 'bcg']
['com', 'eti', 'chi']
[]
['bde']
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Question 2 – Capturing Groups in Regular Expressions 📡

The dataset stored in `data/messy.txt` contains personal information from a fictional website that a user scraped from web server logs. Within this dataset, there are four fields that are of interest to you:
1. Email Addresses 
    - assume they are alphanumeric (dashes do not count) usernames and domain names
2. [Social Security Numbers](https://en.wikipedia.org/wiki/Social_Security_number#Structure)
3. Bitcoin Addresses (alphanumeric strings of long length)
4. Street Addresses

Complete the implementation of the function `extract_personal`, which takes in a string containing the contents of a server log file (like `open('data/messy.txt').read()`) and returns a **tuple of four separate lists** containing values of the 4 pieces of information listed above (in the order listed above). Do **not** keep empty values.

***Note:*** Since this data is messy, your function will be allowed to miss ~5% of the records in each list. Good spot checking using certain useful substrings (e.g. `'@'` for emails) should help assure correctness! Your function will be tested on a sample of the file `messy.txt`.

***Hint:*** There are multiple "delimiters" in use in the file; there are few enough of them that you can safely determine what they are.

```python
# experiment with extract_personal using the file s below
fp = os.path.join('data', 'messy.txt')
s = open(fp, encoding='utf8').read()
```

```python
def extract_personal(s):
    # check for 2-4 characters after the . -- (Done)
    email_address = re.findall(r'[\w]+@[\w]+.\w{2,4}', s)

    # need to clean ssns and btc_address: remove ssn: and bitcoin: from the strings -- Done
    ssns = re.findall(r'ssn:\d{3}-\d{2}-\d{4}', s)
    ssns = [i.strip("ssns:") for i in ssns]

    # check why some outputs returning 'ull' -- (Need to do)
    btc_address = re.findall(r'bitcoin:(?!null)[\w]+', s)
    btc_address = [i.strip("bitcoin:") for i in btc_address]


    # clean \xxx at the end of addresses -- (Done)
    street_address = re.findall(r'\d+ [a-zA-Z ]+', s)
    return email_address, ssns, btc_address, street_address

# 'bitcoin:1BvBMSEYstWetqTFn5Au4m4GFg7xJaNVN2$jk%^3\t,test@test.com,lkj5r%ji|ssn:423-01-9575,530 High Street'

```

```python
# print(extract_personal("bitcoin:1BvBMSEYstWetqTFn5Au4m4GFg7xJaNVN2$jk%^3\t,test@test.com,lkj5r%ji|ssn:423-01-9575,530 High Street"))
# extract_personal('messy.txt')
# print('messy.txt')
```

```python
# don't change this cell, but do run it -- it is needed for the tests
test_fp = os.path.join('data', 'messy.test.txt')
# test_fp = os.path.join('data', 'messy.txt')
test_s = open(test_fp, encoding='utf8').read()
emails, ssn, bitcoin, addresses = extract_personal(test_s)
# emails
# ssn
# bitcoin
# addresses
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Question 3 – TF-IDF 📊

The dataset `data/reviews.txt` contains [Amazon reviews](https://cseweb.ucsd.edu/~jmcauley/datasets.html#amazon_reviews/) for ~200k phones and phone accessories. The dataset has already been "cleaned" for you. In this question, you will complete the implementation of a function, which takes in the reviews dataset as a Series (with one entry per review) as well as a single review, and returns the word that "best summarizes the single review" using TF-IDF.

To do so, implement the two functions below.

#### `tfidf_data`
You will first need to decompose the review into individual words, count their frequencies, and calculate the TF-IDF score for each word. Complete the implementation of the function `tfidf_data`, which takes in the reviews data as a Series (`reviews_ser`) and a single review (`review`) from that Series and returns a DataFrame indexed by the words in `review` with four columns:
- `'cnt'`: the number of times each word is found in the review (telling how frequently a word is used in the `review`)
- `'tf'`: the term frequency for each word (giving an idea of how important each word relative to the other words in the `review`)
- `'idf'`: the inverse document frequency for each word (helping understand how important each word is to the whole `reviews_ser`)
- `'tfidf'` the TF-IDF for each word (measuring how important each word is to `review`, taking into account both the frequency of the word in the `review` and across `reviews_ser`)

You may use a `for`-loop. The words in the outputted DataFrame may appear in any order.

You may also assume that the provided `review` will always be an element of the `reviews_ser` Series, so you do not need to account for division-by-zero errors.

***Hint:*** You may need to use the [`'\b'` character](https://www.regular-expressions.info/wordboundaries.html) somewhere.
    
<br>

#### `relevant_word`

Complete the implementation of the function `tfidf_data`, which takes in the DataFrame that `tfidf_data` returns and returns the word that "best summarizes" the review. If there are multiple "best" summary words, return any one of them.



```python
# experiment with tfidf_data using reviews_ser and review below 
fp = os.path.join('data', 'reviews.txt')
reviews_ser = pd.read_csv(fp, header=None).squeeze("columns")
review = open(os.path.join('data', 'review.txt'), encoding='utf8').read().strip()
```

```python
def tfidf_data(reviews_ser, review):
    def tokenize(string):
        review = re.findall(r'\b\w+\b', string)
        return review
    counts = np.zeros(len(tokenize(review)))
    counts_df = pd.DataFrame().assign(cnt=counts).assign(words=tokenize(review)).groupby("words").count()
    counts_df = counts_df.assign(tf=(counts_df['cnt'] / len(tokenize(review))))
    idf = {}
    for word in counts_df["cnt"].index:
        df = reviews_ser.str.contains(rf'\b{word}\b', case=False).sum()
        idf[word] = np.log(len(reviews_ser) / df)
    counts_df = counts_df.assign(idf=pd.Series(idf))
    counts_df = counts_df.assign(tfidf=counts_df["tf"]*counts_df["idf"])
    return counts_df
```

```python
def relevant_word(out):
    return out["tfidf"].sort_values(ascending=False).index[0]

```

```python
# don't change this cell, but do run it -- it is needed for the tests
fp = os.path.join('data', 'reviews.txt')
reviews_ser = pd.read_csv(fp, header=None).squeeze("columns")
review = open(os.path.join('data', 'review.txt'), encoding='utf8').read().strip()
q3_tfidf = tfidf_data(reviews_ser, review)

try:
    q3_rel = relevant_word(q3_tfidf)
except:
    q3_rel = None
```

```python
q3_tfidf
```

Output:
```
cnt        tf       idf     tfidf
words                                             
5                  1  0.011765  2.271546  0.026724
a                  4  0.047059  0.360392  0.016960
against            1  0.011765  4.266529  0.050194
also               2  0.023529  1.924159  0.045274
and                5  0.058824  0.248188  0.014599
before             1  0.011765  2.820618  0.033184
case               3  0.035294  1.121935  0.039598
chunk              1  0.011765  8.028953  0.094458
color              1  0.011765  2.835119  0.033354
combinations       1  0.011765  7.259820  0.085410
comes              1  0.011765  2.829584  0.033289
cover              2  0.023529  2.682848  0.063126
covers             1  0.011765  3.747910  0.044093
create             1  0.011765  5.814336  0.068404
damage             1  0.011765  4.285037  0.050412
design             1  0.011765  3.055127  0.035943
different          2  0.023529  2.960048  0.069648
from               1  0.011765  1.605962  0.0188
```

```python
q3_rel
```

Output:
```
'chunk'
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## Questions 4 and 5 – Tweet Analysis 🐥

The dataset `data/ira.csv` contains tweets tagged by Twitter as likely being posted by the [Internet Research Agency](https://en.wikipedia.org/wiki/Internet_Research_Agency), the tweet factory facing allegations for attempting to influence US political elections.

Questions 4 and 5 will focus on the following:
- Question 4: Look at the hashtags present in the text and trends in their makeup.
- Question 5: Prepare the dataset for modeling by creating features out of the text fields.

### Question 4 – Hashtags #️⃣

You may assume that a hashtag is any string without whitespace that immediately follows a `'#'` symbol. To clarify, a hashtags can contain a `'#'` symbol, as long as there is no whitespace. For example, `'#data#science'` is considered one individual hashtag, whereas `'#data #science'` is considered two individual hashtags.

#### `hashtag_list`

Complete the implementation of the function `hashtag_list`, which takes in a Series of tweet texts and returns a Series containing a list of hashtags present in each tweet's text. If a tweet's text doesn't contain a hashtag, the Series should contain an empty list for that tweet. Don't include the leading `'#'` symbol in the lists that are returned.

<br>

#### `most_common_hashtag`

Complete the implementation of the function `most_common_hashtag`, which takes in a Series of hashtag lists (as is outputted by `hashtag_list`) and returns a Series consisting of a single hashtag per tweet: 
- If the tweet's text has no hashtags, the entry should in the output Series should be `NaN`.
- If the tweet's text has one distinct hashtag, the entry in the output Series should be that hashtag.
- If the tweet's text has more than one hashtag, the entry in the output Series should be the most common hashtag **in the tweet's text** with respect to **the whole input Series**. If there is a tie for the most common, any of the most common can be returned.
    - For example, if the input Series was `pd.Series([[2], [2], [3, 2, 3]])`, the output would be `pd.Series([2, 2, 2])`. Even though `3` was more common in the third list than `2`, `2` is more common than `3` among all hashtags in the Series.

```python
# The public tests don't test your work on the `ira` data,
# but the hidden tests do.
# So, make sure to thoroughly test your work yourself!
fp = os.path.join('data', 'ira.csv')
ira = pd.read_csv(fp, names=['id', 'name', 'date', 'text'])
```

```python
def hashtag_list(tweet_text):
    hashtags = []
    for line in tweet_text:
        hashtags.append(re.findall(r'#[a-zA-Z0-9_#]+', line))
    return pd.Series([[elem.lstrip("#") for elem in list] for list in hashtags])
```

```python
def most_common_hashtag(tweet_lists):
    all = []
    for lst in tweet_lists:
        all.extend(lst)
    all = pd.Series(all).value_counts()
    to_return = []
    for lst in tweet_lists:
        if len(lst) == 0:
            to_return.append(np.nan)
        elif len(lst) == 1:
            to_return.append(lst[0])
        else:
            dct = {}
            for elem in lst:
                dct.update({elem: pd.Series(all).loc[elem]})
            to_return.append(pd.Series(dct).idxmax())
    return pd.Series(to_return)
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

### Question 5 – Features 📋

Now, create a DataFrame of features from the `ira` data.  That is, create a function `create_features` that takes in a DataFrame `ira` that has just a single column, `'text'`, and returns a DataFrame with the same index as `ira` (i.e. the rows correspond to the same tweets) and the following columns:
* `'num_hashtags'`, the number of hashtags present in the tweet.
* `'mc_hashtags'`, the most common hashtag associated to the tweet (using the result of `most_common_hashtag` from Question 4).
* `'num_tags'`, the number of tags the tweet has.
    - A tag is a string starting with the `'@'` symbol followed by only alphanumerical characters. A`'@'` symbol on its own does not count as a tag.
* `'num_links'`, the number of hyperlinks present in the tweet.
    - A hyperlink is a string starting with `'http://'` or `'https://'`, not followed by whitespaces.
* `'is_retweet'`, a Boolean describing whether the tweet is a retweet. A retweet is a tweet that **begins** with `'RT'`.
* `'text'`, a version of the tweet's text that is cleaned according to the following steps, **in this exact order**:
    1. All meta-information above (retweet info, tags, hyperlinks, and hashtags) should be replaced with a single space.
    2. Everything other than letters, numbers, and spaces should be replaced with a single space.
    3. All letters should be lowercase.
    4. All words should be separated by exactly one space, and leading/trailing whitespace should be removed (stripped).
    
The columns in the outputted DataFrame must be in the order `['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']`. (Remember, the DataFrame that `create_features` is called on only has a single column, `'text'`.)

***Notes:***
- It's a good idea to make helper function for each column.
- The `\w` character class in regex **does not** refer to letters, numbers, and spaces (or even just letters and numbers). As such, you can't use it here!
- `create_features` will take a while to run on the entire dataset – test it on a small sample first!
- A `'@'` by itself is not a tag - do not account for empty strings as tags

```python
# The doctests/public tests don't test your work on the `ira` data,
# but the hidden tests do.
# So, make sure to thoroughly test your work yourself!
fp = os.path.join('data', 'ira.csv')
ira = pd.read_csv(fp, names=['id', 'name', 'date', 'text'])
```

```python
# don't change this cell, but do run it -- it is needed for the tests
# (yes, we know it says "hidden" – there are still truly hidden tests in this question)
fp_hidden = 'data/ira_test.csv'
ira_hidden = pd.read_csv(fp_hidden, header=None)
text_hidden = ira_hidden.iloc[:, -1:]
text_hidden.columns = ['text']
test_hidden = create_features(text_hidden)
```

```python
def create_features(ira):
    def clean_text(series):
        cleaned = []
        for text in series:
            text = re.sub(r'https?://\S+', ' ', text)
            text = re.sub(r'^RT', ' ', text)
            text = re.sub(r'@[a-zA-Z0-9_]+', ' ', text)
            text = re.sub(r'#[a-zA-Z0-9_]+', ' ', text)
            text = re.sub(r'[^a-zA-Z0-9 ]', ' ', text)
            text = text.lower()
            text = re.sub(r'\s+', ' ', text).strip()
            cleaned.append(text)
        return pd.Series(cleaned)

    mod = ira.copy()
    count_hashtags = hashtag_list(ira["text"]).apply(len)
    mod = mod.assign(num_hashtags=count_hashtags)
    mod = mod.assign(mc_hashtags=most_common_hashtag(hashtag_list(ira["text"])))

    ats = []
    for i in ira["text"]:
        ats.append(len(re.findall(r'@[a-zA-Z0-9]+', i)))
    mod = mod.assign(num_tags=ats)
    links = []
    for i in ira["text"]:
        links.append(len(re.findall(r'(https://[\w]+)|(http://[\w]+)', i)))
    mod = mod.assign(num_links=links)
    retweet = []
    for i in ira["text"]:
        if len(re.findall(r'(^RT)', i)) >= 1:
            res = True
        if len(re.findall(r'(^RT)', i)) == 0:
            res = False
        retweet.append(res)
    mod = mod.assign(is_retweet=retweet)
    return mod.assign(text=clean_text(ira["text"]))

create_features(text_hidden)
```

Output:
```
text  num_hashtags  \
0        the best exercise to lose belly fat in 2 weeks             4   
1     no more reserved parking spaces for journalist...             0   
2     not trumpriots but don t let the left frame th...             1   
3     was interrupted by you re a rapist heckle at c...             2   
4           billium payium pretty much self explanatory             1   
...                                                 ...           ...   
996   warriors rival spurs highlight nba s internati...             1   
997                         now playing juicy j trap gt             0   
998              2016 fourth of july fireworks displays             0   
999                                  absolutely nothing             3   
1000  house committee launches review of fda crimina...             1   

              mc_hashtags  num_tags  num_links  is_retweet  
0                Exercise         0          2       False  
1                     NaN         1          1      
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

```python
testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
test = pd.DataFrame(testdata, columns=['text'])
out = create_features(test)
anscols = ['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']
ansdata = [['text cleaning is cool', 3, 'NLP1', 1, 1, True]]
ans = pd.DataFrame(ansdata, columns=anscols)
# bool((out == ans).all().all())
print(out)
```

Output:
```
text  num_hashtags mc_hashtags  num_tags  num_links  \
0  text cleaning is cool             3        NLP1         1          1   

   is_retweet  
0        True
```

## Congratulations! You're done Lab 7! 🏁

As a reminder, all of the work you want to submit needs to be in `lab.py`.

To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.

Once you've finished the lab, you should open the command line and run, in the directory for this lab:

```
python lab-validation.py
```

**This will run all of the `grader.check` cells that you see in this notebook, but only using the code in `lab.py` – that is, it doesn't look at any of the code in this notebook. If all of your `grader.check` cells pass in this notebook but not all of them pass in your command line with the above command, then you likely have code in your notebook that isn't in your `lab.py`!**

You can also use `lab-validation.py` to test individual questions. For instance,

```
python lab-validation.py q1 q2 q4
```

will run the `grader.check` cells for Questions 1, 2, and 4 – again, only using the code in `lab.py`. [This video](https://www.loom.com/share/0ea254b85b2745e59322b5e5a8692e91?sid=5acc92e6-0dfe-4555-9b6a-8115b6a52f99) how to use the script as well.

Once `python lab-validation.py` shows that you're passing all test cases, you're ready to submit your `lab.py` (and only your `lab.py`) to Gradescope. Once submitting to Gradescope, make sure to stick around until all test cases pass.

There is also a call to `grader.check_all()` below in _this_ notebook, but make sure to also follow the steps above.

```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!

q5 results: All test cases passed!
```

Tests for q1 (points: None)
```
>>> import doctest
>>> doctest.run_docstring_examples(match_1, {'match_1': match_1})
>>> doctest.run_docstring_examples(match_2, {'match_2': match_2})
>>> doctest.run_docstring_examples(match_3, {'match_3': match_3})
>>> doctest.run_docstring_examples(match_4, {'match_4': match_4})
>>> doctest.run_docstring_examples(match_5, {'match_5': match_5})
>>> doctest.run_docstring_examples(match_6, {'match_6': match_6})
>>> doctest.run_docstring_examples(match_7, {'match_7': match_7})
>>> doctest.run_docstring_examples(match_8, {'match_8': match_8})
>>> doctest.run_docstring_examples(match_9, {'match_9': match_9})
>>> doctest.run_docstring_examples(match_10, {'match_10': match_10})
```
failure_message: doctests


Tests for q2 (points: None)
```
>>> emails[0] == 'test@test.com'
True
```
```
>>> ssn[0] == '423-01-9575'
True
```
```
>>> bitcoin[0] == '1BvBMSEYstWetqTFn5Au4m4GFg7xJaNVN2'
True
```
```
>>> addresses[0] == '530 High Street'
True
```


Tests for q3 (points: None)
```
>>> bool(q3_tfidf['cnt'].sum() == 85)
True
```
```
>>> 'before' in q3_tfidf.index
True
```
```
>>> bool(q3_tfidf.loc['a', 'cnt'] == 4)
True
```
```
>>> bool(np.isclose(q3_tfidf.loc['phone', 'idf'], 0.820703, atol=0.5))
True
```
```
>>> q3_rel in q3_tfidf.index
True
```


Tests for q4 (points: None)
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = pd.DataFrame(testdata, columns=['text'])['text']
>>> out = hashtag_list(test)
>>> out.iloc[0] == ['NLP', 'NLP1', 'NLP1']
True
```
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = hashtag_list(pd.DataFrame(testdata, columns=['text'])['text'])
>>> most_common_hashtag(test).iloc[0] == 'NLP1'
True
```
```
>>> out = pd.Series([[2], [2], [3, 2, 3]])
>>> bool((most_common_hashtag(out) == 2).all())
True
```


Tests for q5 (points: None)
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = pd.DataFrame(testdata, columns=['text'])
>>> out = create_features(test)
>>> anscols = ['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']
>>> ansdata = [['text cleaning is cool', 3, 'NLP1', 1, 1, True]]
>>> ans = pd.DataFrame(ansdata, columns=anscols)
>>> bool((out == ans).all().all())
True
```
```
>>> test_hidden.columns.tolist() == ['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']
True
```
```
>>> bool(400 <= (test_hidden['num_tags'] == 0).sum() <= 450)
True
```
```
>>> bool(450 <= (test_hidden['num_links'] == 0).sum() <= 550)
True
```
